# Exercise 3 - Levenshtein distance
(3 points)

Finish the `calcLevenshteinDistance` method by implementing the calculation of the Levenshtein distance of the two given Strings. Use dynamic programming as described in the lecture slides.

#### Hints

- Please note that implementations that do not use dynamic programming (e.g., a recursive solution) may cause either a timeout or a stack overflow since your solution will be tested with long strings. In the provided tests, you will see a test case with two very long strings. Your solution is expected to handle them within 100ms. Note that you still can achieve points if your solution is slower, but it should be seen as a warning that it might cause issues during the evaluation.

#### Notes

- You are free to use a different IDE to develop your solution. However, you have to copy the solution into this notebook to submit it.
- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Known issues
  - All global variables will be set to void after an import.
  - Missing spaces arround `%` (Modulo) can cause unexpected errors so please make sure that you have added spaces around every `%` character.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Make sure that all necessary imports are listed at the beginning of your cell.
  - Run a final check of your solution by
    - click on _restart the kernel, then re-run the whole notebook_ (the fast forward arrow in the tool bar)
    - wait fo the kernel to restart and execute all cells (all executable cells should have numbers in front of them instead of a `[*]`) 
    - Check all executed cells for errors. If an exception is thrown, please check your code. Note that although the error might look cryptic, until now we never encounter that an exception was caused without a valid reason inside of the submitted code. A good way to check the code is to copy the solution into a new class in your favorite IDE and check
      - errors reported by the IDE
      - imports the IDE adds to your code which might be missing in your submission.
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
public int calcLevenshteinDistance(String string1, String string2) {
    int distance = 0;
		// YOUR CODE HERE

		int[] changes = new int[string2.length() + 1];
		int index2 = 0;
		while(index2<changes.length){

			changes[index2] = index2;
			index2++;
		}
        int index1 = 1;
       while(index1<=string1.length()){
			changes[0] = index1;
			int previousEdit = index1 - 1;
			for ( index2 = 1; index2 <= string2.length(); index2++) {

				// whenever we found a change in the character, we add 1 in the
				// previous editdistance and then run the minimum formula on it
				// with the other distance
				int newEditDistance = Math.min(1 + Math.min(changes[index2 - 1], changes[index2]),
						string1.charAt(index1 - 1) == string2.charAt(index2 - 1) ? previousEdit : previousEdit + 1);
				previousEdit = changes[index2];
				changes[index2] = newEditDistance;
			}
                index1++;
		}
		distance = changes[string2.length()];
		return distance;
}

// This line should make sure that compile errors are directly identified when executing this cell
// (the line itself does not produce any meaningful result)
Arrays.sort(new int[calcLevenshteinDistance("a","a")]);
```

# Evaluation

- Run the following cell to test your implementation.
- You can ignore the cells afterwards.

<!-- grade_id: cell-autograde-visible, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

public void checkLevenshteinDistance(String s1, String s2, int expectedDistance) {
    checkLevenshteinDistance(s1, s2, expectedDistance, 1);
}
    
public void checkLevenshteinDistance(String s1, String s2, int expectedDistance, long maxRuntime) {
    try {
        long time = System.currentTimeMillis();
        int result = calcLevenshteinDistance(s1, s2);
        time = System.currentTimeMillis() - time;
        Assertions.assertEquals(expectedDistance, result, "For levenshtein(\"" + s1 + "\",\"" + s2
                + "\") your solution returned " + result + " while " + expectedDistance + " has been expected.");
        if(time > maxRuntime) {
            System.out.println("Test successfully completed. Your solution took more time (" 
                               + time + "ms) than expected (" + maxRuntime + "ms).");
        } else {
            System.out.println("Test successfully completed.");
        }
    } catch (AssertionFailedError e) {
        throw e;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error:");
        throw e;
    }
}

checkLevenshteinDistance("abababab", "aaaaaaaa", 4);
checkLevenshteinDistance("rupert", "Robert", 3);
checkLevenshteinDistance("Robert", "rupert", 3);
checkLevenshteinDistance("Robert", "Robert", 0);
checkLevenshteinDistance("abcde", "", 5);
checkLevenshteinDistance("", "abcde", 5);


String word1 = "bqsvtzfhktnqovnjmkktluayihcfelhfqfsordvkxrmweenobehobnokrzabxcglfavxwyqlhjjlfxyx"
    + "ddptklfuwvtvwgecpdskvovhrrlaoqzjqkpauhpkkrxtrwmrevimykuyvstfgsmeegqjvxcstegiiqgl"
    + "auprgfuchymgvznavfpxnflxlujvtsshwhsnkunetfiwvqppaanugcovhnvwwedxywtkswmfeccatpna"
    + "fzntnagkjpsinosekeahzmqtcpjkaxzmunclbcurppwumfnmxlrfgjxsddblhiorwiaqcjlbnbfkdjfr"
    + "fqeqzrlfoiejfcmwdrcxneoquogbbuwdsykxvjpmfahmoybrqmneyrlutiunrbhfclzzhdtlxnwivoeg"
    + "zfzxmargbkxkjnqiuqsgakhzgxkxtchewatusacyeodaankidwezlopjabyzaxpxzffxnqaredklqron"
    + "slyzqgtfdjkaiwdcgalwvgpnlbxvuyqmhxsudcthfgziahugmjdhpqjucyvffviupnzmdtbatqbfewhy"
    + "jyfsijxxuyoyonldvclzrdoujbaskietubkfoopbeucqrgxzvmkwcdmbuglfcnkizkefskvmqazbtyyn"
    + "lrwptebsqvrtwtfukoesbjtgcxesvimzzdyoqvnclbedqhhgirhdyexourkwhzpjxjctdvqpiprfedsr"
    + "xywdzncqihboggtekrknjbrlflrfegxjbvaqqjcodbzmvgefbateguibcajueuohwftimarzgjjygyng"
    + "ayihcjnlnhidcooojummtbjuaddjsfyxqqpmphsuicxeeaizqrsrdomxhxwymjtfuigiijwowdoolbdv"
    + "bvwtipnybxndkubuzskeycpjuuiutokzlibbprywcnovjrbrhmlsoyumyxuaswbkoopaxbndrjyocqle"
    + "ishdwuuummzmtovjidjetgsfftrfmiqjckfphsvt";
String word2 = "bqsvtzfhktnuosnjmkktluayihcfelhfqfsordvkxrmweenobelobnocrzabxcglnavxcyqlhjjlfxyx"
    + "odptklfurvtvwiecpdrkvovhrrlnoizjqkpauhpkkrxtrwmrmvimykuyvstfgsmeenqjvxcstegiiqgl"
    + "aufrgfuchymgvznavfpxnflxlujvtssxwhsnkuneqfiwvqpdaanugcnvhavwwedxywtkswmfeccatpna"
    + "fzntnagkjpsinosekeahzmqtcpnkapzmunclbcurppwumfnmxmrfglxjddblhiorwiaqcjlbnbfkdjfw"
    + "fqeqzrjfoiejkcmwdrcxneonuogbmtwdsfkxvjpkfahmrynrqgneyrfutiunrbhfclzzhdtlxpwinocg"
    + "zfzxmargbkzkjnqiuqsgakfzgxkctbwewatusacyeodaankidwyzlopqabyzaxpxzffxnqaredkjqron"
    + "slyzqgtfdokaiwdcgalwvtpnlbxvuyqmhxsudcthfgzidhugmjdhpqjucyvffviupnzbdnbatqbfewhy"
    + "jyflcjxxuyoyonldvcpardoujbaskietubkfoopbeucqrgxzvmkwcdmbuglfqnkizkefhkvmqazbtkvn"
    + "lrwptebsqvrtwtfukoesbjtgcxesvimzzdpoqvnclbedqhomcrhdyexouukwwzpfxjctdiqpiprfedsr"
    + "xyodzncqijbogutekrknbbrlflrfegxjbvaqqjcodbzmvgefbateguiscajueuohjftimvrzgjjygymg"
    + "zyihcjnlnzidcaoojulmtbjuaddjslyxqlpmxhsuibxeeaizqrkrdomxhxwymjtfuigiijwowdoolbdv"
    + "svwtihnybxndkubhzskeycpjuuiutokzlibbprywcnovjrbrvllsogumyxuwswqkoouaxbnidjywcqle"
    + "ishbwuuummzmtovjidjetgsfftrfmiqjckfphsvt";
checkLevenshteinDistance(word1, word2, 99, 100); // ~100ms are still ok
checkLevenshteinDistance(word2, word1, 99, 100);
```

Output:
```
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
```

<!-- grade_id: cell-autograde-1, points: 2, grade: true, solution: false -->
```java
// Ignore this cell
```

<!-- grade_id: cell-autograde-2, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
```

```java

```