# Exercise 1 - Preprocessing
(3 points)

Given a text, your implementation of the `preprocess` method should perform a preprocessing generating a list of tokens. This preprocessing should include 1) lowercasing, 2) sentence splitting and 3) the removal of punctuations.

For the input:
```
She said:"I know that she likes English food!"
```
the preprocessing should generate a `List<String>` containing
```
[<s>, she, said, i, know, that, she, likes, english, food, </s>]
```
(The line above is the output of `List.toString()`, i.e., the squared brackets and commas are not part of the tokens.)

Your preprocessing should:
* mark the beginning of a sentence with `<s>` and its end with `</s>`. For simplicity, you can assume that every occurence of `.`, `!` or `?` will end a sentence.
* remove all non alphanumeric characters, i.e., the tokens should contain only lowercased, alphabetical character or digits. It is known that this will lead to some faulty behavior (e.g., `won't` will be cut into two words `won` and `t`). However, to keep the exercises easy, we will ignore these special cases.
* not generate empty tokens or an empty sentence, i.e., `<s>, </s>` should never occur (while `</s>, <s>` is of cause necessary).

#### Notes

- You are free to use a different IDE to develop your solution. However, you have to copy the solution into this notebook to submit it.
- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Known issues
  - All global variables will be set to void after an import.
  - Missing spaces arround `%` (Modulo) can cause unexpected errors so please make sure that you have added spaces around every `%` character.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Make sure that all necessary imports are listed at the beginning of your cell.
  - Run a final check of your solution by
    - click on _restart the kernel, then re-run the whole notebook_ (the fast forward arrow in the tool bar)
    - wait fo the kernel to restart and execute all cells (all executable cells should have numbers in front of them instead of a `[*]`) 
    - Check all executed cells for errors. If an exception is thrown, please check your code. Note that although the error might look cryptic, until now we never encounter that an exception was caused without a valid reason inside of the submitted code. A good way to check the code is to copy the solution into a new class in your favorite IDE and check
      - errors reported by the IDE
      - imports the IDE adds to your code which might be missing in your submission.
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
/**
 * The token used to indicate the beginning of a new sentence.
 */
public static final String SENTENCE_START = "<s>";
/**
 * The token used to indicate the end of a sentence.
 */
public static final String SENTENCE_END = "</s>";

public static List<String> preprocess(String text) {

		List<String> tokenList = new ArrayList<>();
		String text_array[] = text.replaceAll("[^ a-zA-Z0-9.?!]"," ").replaceAll("[.!?]", " </s> <s> ").toLowerCase().split(" ");
		tokenList.add("<s>");
		for(String token: text_array){
			if(!(token.equals("")))
				tokenList.add(token);
		}
		tokenList.remove(tokenList.size()-1);
		return tokenList;
	}
// This line should make sure that compile errors are directly identified when executing this cell
// (the line itself does not produce any meaningful result)
preprocess("").clear();
```

# Evaluation

- Run the following cell to test your implementation.
- You can ignore the cells afterwards.

<!-- grade_id: cell-autograde-visible-1, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

public static void checkPreprocessingResult(String text, String expectedTokens[]) throws Exception {
    try { 
        List<String> preProcResult = preprocess(text);
        Assertions.assertNotNull(preProcResult, "The given preprocessing result is null.");
        String[] result = preProcResult.toArray(new String[preProcResult.size()]);
        Assertions.assertArrayEquals(expectedTokens, result, 
                                     "Your preprocessing generated " + 
                                     Arrays.toString(result) + 
                                     " but the expected result was " + Arrays.toString(expectedTokens));
        Arrays.toString(expectedTokens);
        System.out.println("Test successfully completed.");
    } catch (AssertionFailedError e) {
        throw e;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error:");
        throw e;
    }
}

String text1 = "She said:\"I know that she likes English food!\"";
String[] expectedTokens1 = new String[]{"<s>", "she", "said", "i", "know", "that", "she", "likes",
                                        "english", "food", "</s>"};
checkPreprocessingResult(text1, expectedTokens1);

text1 = "I do not eat this soup. No, I will not eat this soup. Soup is not what I like to eat 24/7.";
expectedTokens1 = new String[]{"<s>", "i", "do", "not", "eat", "this", "soup", "</s>", "<s>",
                               "no", "i", "will", "not", "eat", "this", "soup", "</s>", "<s>",
                               "soup", "is", "not", "what","i", "like","to", "eat", "24", "7",
                               "</s>"};
checkPreprocessingResult(text1, expectedTokens1);
```

Output:
```
Test successfully completed.
Test successfully completed.
```

<!-- grade_id: cell-autograde-1, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
```

<!-- grade_id: cell-autograde-2, points: 2, grade: true, solution: false -->
```java
// Ignore this cell
```